***DISCLAIMER (Read this carefully)***

Before you turn this assignment in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel
Restart) and then run all cells (in the menubar, select Cell
Run All). Do NOT add any cells to the notebook!

Do not forget to submit both the notebook AND the files in the data/ subfolder according to the CoC!
Make sure you fill in any place that says YOUR CODE HERE or YOUR ANSWER HERE , as well as your name and group below:

# Assignment 2 (Group)
When carrying out a Data-Science project, screening and selecting appropriate data sources for the tasks at hand comes at the beginning. This assignment is about accessing and characterising potential data sources in teams of three. The teams have been randomly assigned. BEWARE! In Assignment 5, you will be asked to provide answers to those questions. Make sure that combining the two datasets makes sense from an analytical perspective!

-----
## Step 0 (2 points)

Find two data sets online (from one or several sources) that would be interesting to combine and create ***data citations*** as Python dictionaries. 

The data sets should fulfill the following requirements:

* Each data set must have a different file format (either CSV, XML, or JSON), please choose
  - one CSV file (dataset1) 
  - and one JSON or XML file (dataset2)

* The two datasets should not be two variations of each other (i.e. simply the same dataset for two different regions or timeframes or from the same source just in two different formats)
* Workable data-set sizes: The selected or extracted data sets should have thousands of entries (>= 1000), but not more than (<=) 10000 entries. Be "entries we mean rows or distinguishable key-value pairs). If larger, use an excerpt from the original data set. Justify in detail the extraction criteria in the markdown cell below and 
  1) add the code used for the extraction in the code cell or describe how you filtered the sample  
  2) make the extracted dataset also available at a downloadable URL (for instance in a Github repository, [here](https://raw.githubusercontent.com/AxelPolleres/simple_dataset_sharing_repo/main/test.csv)'s an example)
  3) name the new `resourceURL` in the data citation.
* You may start from (but you are not limited to) the resource collections hinted at [in the Unit 2 slides](https://datascience.ai.wu.ac.at/ws21/dataprocessing1/unit2.html#slide-53).

* Important: The use of datasets from kaggle.com and other curated collections of datasets with accompanying tutorials on processing and analysis (as highlighted to you in Unit 2) is **discouraged**. You are required to use **primary data sources**: This is mainly because we want you to work on data sets that have not been processed with some analysis in mind, so that you show that you can handle (messy) data sets harvested on the brownfields of Data Science. Besides, such curated datasets have been repeatedly used in ready-made case and tutorial work, which makes it basically impossible for us to establish whether your submissions are genuine contributions of yours. There is one viable option: Work backwards from the Kaggle data set to the original data source, obtain updated data from there, and start from there.


* Please adhere to the CoC.

[Data citations](http://blogs.nature.com/scientificdata/2016/07/14/data-citations-at-scientific-data/) must contain the following details:
- creator: provider organisation / author(s) of the data set, e.g. "Zentralanstalt für Meteorologie und Geodynamik (ZAMG)"
- catalogName: Names of the data repository and/or the Open Data portal used, e.g. Open Data Österreich"
- catalogURL: URL of th repository / portal, e.g. "https://www.data.gv.at/"
- datasetID: (specific to the data repository), e.g. "https://www.data.gv.at/katalog/dataset/zamg_meteorologischemessdatenderzamg"
- resourceURL: a URL where the CSV, XML or JSON file can be downloaded, e.g. "https://www.football-data.co.uk/new/JPN.csv"
- pubYear: Dataset publication year, i.e. since when it is published, e.g. "2012"
- lastAccessed: when have you last accessed the dataset (i.e. datetime of accessing, obtaining a copy of the data set) in ISO Format? e.g. "2021-03-08T13:55:00"

One final note: as mentioned above, if you want to use a repository for your file download (e.g. github), you are allowed to do that. The most important part is that the URL can be accessed stably for each dataset you have chosen. 

Store the data citation in a dictionary for each of the datasets:

<!-- grade_id: cell-cea2669d70d8d76a, points: 0, grade: true, solution: true -->
```python
dataset1= {
    "creator" : "U.S. Department of Health & Human Services" ,
    "catalogName" : "US Data Catalog" ,
    "catalogURL" : "catalog.data.gov.us" ,
    "datasetID" : "5320bb9d-7476-4408-a410-e6ca30f13b3c" ,
    "resourceURL" : "https://raw.githubusercontent.com/Zedtisfying/DP1-Set-1/main/Indicators_of_Anxiety_and_Depression.json"  ,
    "pubYear" : "2020"  ,
    "lastAccessed" : "2024-04-16T12:29:36",
}

dataset2= {
    "creator" : "U.S. Department of Health & Human Services" ,
    "catalogName" : "US Data Catalog" ,
    "catalogURL" : "catalog.data.gov.us" ,
    "datasetID" : "732f92fd-4775-44c9-a228-3cd2121abe82" ,
    "resourceURL" : "https://raw.githubusercontent.com/Zedtisfying/DP1-Set-1/main/Indicators_of_Health_Insurance_Coverage_at_the_Time_of_Interview_(Last).csv"  ,
    "pubYear" : "2020"  ,
    "lastAccessed" : "2024-04-16T12:30:58",
}
```

<!-- grade_id: cell-d3cc293b6207d1c7, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_in, assert_true
import traceback
import sys
import os

assert_equal(type(dataset1), dict)
assert_equal(type(dataset2), dict)

```

Use the following structure for your answer below:

Data set 1

(Describe the source and the general content of the dataset and why you chose it)

Data set 2

(Describe the source and the general content of the dataset and why you chose it)

Project ideas

(Describe in your own words, which kind of tasks could be addressed by combining the selected data sets, esp. how the two data sets fit together and what complementary information they contain; Formulate a question that could be potentially answered by combining data from both datasets; how could the data sets be combined exactly? 250 words max. BEWARE! In Assignment 5, you will be asked to provide answers to those questions. Make sure that combining the two datasets makes sense from an analytical perspective!)

<!-- grade_id: cell-9bc09f21e0c42050, points: 1, grade: true, solution: true -->
Data set 1

During the initial phases of the pandemic, the U.S. Census Bureau, in partnership with five federal agencies, launched the Household Pulse Survey to gather data on the social and economic effects of Covid-19. This JSON file provides insights into the mental impacts, specifically depression and anxiety, categorized by state and documented over various timeframes.The provided file has been shortened to fit the task requirements, and had the statistical values removed to keep the most relevant parts.

Data set 2

The second dataset consists of the results of a survey, in which the applicants were asked their insurance state. This CSV file consists of the answers to this survey documented over different timeframes. The file has been shortened to cover only specific time frame to match with the first dataset additionally the columns including the calculations of statistical values and intervalls have been removed.

Project ideas: This project aims to investigate the influence of insurance coverage on the severity of anxiety symptoms among residents across different states during the Covid-19 pandemic. By merging two datasets—one detailing anxiety severity by state and timeframe and the other illustrating insurance status categorized by numbers within the same timeframe—this study seeks to uncover potential correlations between mental health outcomes and access to healthcare resources at a regional level. These two datasets will be combined based on the state variables to examine whether certain states exhibit higher levels of anxiety symptoms and whether residents' insurance coverage varies by state. The inclusion of multiple timeframes allows for comparisons of anxiety levels within the same states over different periods. With the provided datasets we aim to answer how the insurance coverage impacts the severity of anxiety symptoms among residents of different states during the Covid-19 pandemic.

------
## Step 1 - File Access (3 points)

Write a Python function `accessData` that takes the dataset dictionary created in step 0 as an input and returns an extended dictionary including following additions:

* Write code that accesses the dataset from its `resourceURL` using the python `requests` package:
 * detects whether it's and XML, CSV or JSON file by
     * checking whether the download URL **ends** with suffix "xml", "json", "csv" 
     * checking whether the "Content-Type" HTTP header field contains information about the format, hinting on XML, JSON or CSV, i.e., check whether the substring XML, JSON or CSV appears in the "Content-Type" header in either upper- or lowercase. 
 * Detects the file size from the HTTP header (converted to KB) of each data set, clearly documenting your actions (e.g. through commented code).

The result of the code below should extend your dictionaries `dataset1` and `dataset2` with two keys named 
* `"detectedFormat"` (which has one of the following values: `"XML"`, `"JSON"`, `"CSV"`, or `"unknown"`, if nothing could be detected from checking the suffix or HTTP header, or if the information in both was inconsistent)
* and `"filesizeKB"` which contains the filesize in KB (Conversion should be done accordingly to decimal SI prefixes) from the number of bytes in the header-information. If there is no respective header information return 0.
* If the detected format is `"unknown"`, the expected filesize to be returned is also 0


<!-- grade_id: cell-87173edcb1445261, points: 0, grade: false, solution: true -->
```python
import requests

def accessData(datadict):
    # YOUR CODE HERE
    URL = datadict["resourceURL"]
    
# check if URL ENDS with suffix 
    # split URL by dots so that suffix is at index -1 
    suffix = URL.split(".")
    
    
    # what do they mean by "ends with suffix"? only .json or also .json?accessType=DOWNLOAD
    for ftype in ["csv", "json", "xml"]:
        if ftype in suffix[-1]:
            detectedFormatURL = ftype
            break
        else: detectedFormatURL = "unknown"

# checks header information
    r = requests.head(URL)
    rhead = r.headers
    
    # "detectedFormat" (which has one of the following values: "XML", "JSON", "CSV", or "unknown", if nothing could be detected from checking the suffix or HTTP header, or if the information in both was inconsistent)
    # "if the information is both was inconsistent"
    # is it also inconsistent when we dont get a value for content-type? (key content-type not available)
    conttype = rhead["content-type"].lower() if "content-type" in rhead else "unknown"  #detectedFormatURL <- davor # else unknown <- Jonas
    contleng = int(rhead["content-length"])/1000 if "content-length" in rhead else 0
    
   
    detectedFormatHeader = ""
    # checks if file name strings are substrings of the "content-type" header value 
    for ftype in ["csv", "json", "xml"]:
        if ftype in conttype:
            detectedFormatHeader = ftype
            break
        else: detectedFormatHeader = "unknown"
    
    valid_formats = ["csv", "json", "xml"]
    
    # checks if the two file formats found in the suffix and the header information are the same and creates new dictionary pairs
    ## 4 possibilities
        ## if both are the same and in valid_formats return that
        ## One is in valid_formats and one is unknown return the one in validFormats
        ## both unknown --> unknown
        ## both are in valid_formats but not the same --> unknown
    if detectedFormatURL in valid_formats and detectedFormatURL == detectedFormatHeader:
        datadict.update({"detectedFormat": detectedFormatURL.upper(), "filesizeKB": contleng})
    
    elif detectedFormatURL == "unknown" and detectedFormatHeader == "unknown":
        datadict.update({"detectedFormat": "unknown", "filesizeKB": 0})
        
    elif detectedFormatURL in valid_formats and detectedFormatHeader in valid_formats and detectedFormatURL != detectedFormatHeader:
        datadict.update({"detectedFormat": "unknown", "filesizeKB": 0})
    
    elif detectedFormatURL in valid_formats:
        datadict.update({"detectedFormat": detectedFormatURL.upper(), "filesizeKB": contleng})
    
    elif detectedFormatHeader in valid_formats:
        datadict.update({"detectedFormat": detectedFormatHeader.upper(), "filesizeKB": contleng})
   
    
   
    return datadict
    raise NotImplementedError()
```

<!-- grade_id: cell-09b529ecf9606ac6, points: 0.24, grade: true, solution: false -->
```python
# Basic tests to see if your solution meets the foundational demands described in the task description
from nose.tools import assert_equal, assert_in, assert_true
dataset1= accessData(dataset1)
dataset2= accessData(dataset2)
assert_in(dataset1["detectedFormat"], ["XML", "JSON", "CSV", "unknown"])
assert_in(dataset2["detectedFormat"], ["XML", "JSON", "CSV", "unknown"])
assert_true(isinstance(dataset1["filesizeKB"], (int, float)))
assert_true(isinstance(dataset2["filesizeKB"], (int, float)))
```

<!-- grade_id: cell-b1ef1760abd62f3f, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-b16c0da0f2b8dc22, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-e19b418cde2ba027, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-278bfc1d2c26772d, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-1fb00904a60db5c7, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-affc2fa2a6b001c9, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-00e4aadd3923d75d, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-d0566ccfef366cdd, points: 0.22, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

Please explain your findings, using the following structure for your answer below (in "other remarks" you can explain, for instance, why you think your code did not detect the correct format, if needed)

Data set 1

(format, size, other remarks)

Data set 2

(format, size, other remarks)

<!-- grade_id: cell-40248ad63aa2baea, points: 1, grade: true, solution: true -->
Data set 1 

CSV file, 65.646kB, created by U.S. Department of Health & Human Services, from US Data Catalog(catalog.data.gov.us), datasetID 732f92fd-4775-44c9-a228-3cd2121abe82, URL https://raw.githubusercontent.com/Zedtisfying/DP1-Set-1/main/asdfasdfasdfasdfasdf.csv, published in 2020 last accessed on 2024-04-16T12:29:36.

Data set 2

JSON file, 12.499kB,created by U.S. Department of Health & Human Services, from US Data Catalog(catalog.data.gov.us), datasetID 5320bb9d-7476-4408-a410-e6ca30f13b3c URL https://raw.githubusercontent.com/Zedtisfying/DP1-Set-1/main/Indicators_of_Anxiety_and_Depression.json, published in 2020 last accessed on 2024-04-16T12:30:58.

-----
## Step 2  (5 points) - Format Validation

Establish that the two data files obtained are well-formed according to the detected data format (CSV, JSON, or XML). That is, the syntax used is valid according to accepted syntax definitions. Are there any violations of well-formedness?


Proceed as follows (for each data file, in turn): according to the "suspected" data format from Step 1:

  1. Use an _online validator_ for CSV, XML, and JSON, respectively, to confirm whether the files you downloaded in Step 1 are well-formed for the respective file format, document your findings and modify the file as described: 

   a. **Case 1**: no well-formedness errors were detected: 
    * Generally describe at least 3 well-formedness checks that your data sets, depending on its "suspected" format (against the background knowledge of Unit 2) should fulfill;
    * Store a local copy of the file called `data_notebook-[notebook-nr.]_[name].[file extension]` in the `data/` subfolder
    * Create another local copy of your data file called `data_notebook-[notebook-nr.]_[name]-invalid.[file extension]` and introduce a selected well-formedness violation (one occurrence) therein;
    * document that the online validator you used finds the error you introduced

   b. **Case 2**: well-formedness errors occurred:
    * Document the occurrences by printing out the error message and describe the types of well-formedness violation that were reported to you.
    * Store a local copy called `data_notebook-[notebook-nr.]_[name]-invalid.[file extension]`  in the `data/ subfolder`
    * Create another local copy called `data_notebook-[notebook-nr.]_[name].[file extension]`, of your data file that fixes the well-formedness violations therein manually.  
    
**Please note that the datasets in the `data/` subfolder are for documentation only. Do not access those for subsequent steps!**
    

  2. Write a Python function `parseFile(datadict, format)` that that accesses the dataset from its `resourceURL`. The dataset should then be checked accordingly the given parser for the parameter `format` to check the following:
     * CSV: Returns `True`, if a consistent delimiter out of `",",";","\t"` can be detected, such that each row has the same (> 1) number of elements, otherwise False
     * JSON: Returns `True` if the file can be parsed with the `json` package, catching any parsing exceptions.
     * XML: Returns `True` if the file can be parsed with the `xmltodict` package, catching any parsing exceptions.
     * Returns `False` if any other format is supplied by the parameter.
     
In order to handle parsing exceptions and errors from the used packages, you can use [catching exceptions](https://docs.python.org/3/tutorial/errors.html), such that the program does not simply fail to check whether the file is parseable as the format specified in `format`    

Use the following structure for your answer in the cell below to document **Step 2.1**:

***Data set 1***

*(validator used, validation results, describe the modification to fix the file or to create an invalid version of it)*

***Data set 2***

*(validator used, validation results, describe the modification to fix the file or to create an invalid version of it)*


<!-- grade_id: cell-7326ff597819ce15, points: 2, grade: true, solution: true -->
Data set 1

online validator,https://jsonchecker.com/,

error message:

"Parse error on line 1: Indicator,Group,Stat ^ Expecting 'STRING', 'NUMBER', 'NULL', 'TRUE', 'FALSE', '{', '[', got 'undefined'"

The description for the fifth column was missing. It has been replaced with the corresponding column indicator. Additionally, one row containing invalid data types for certain columns, along with empty cells in other columns, has been removed from the dataset for data integrity purposes.

Data set 2

online validator,https://onlinetools.com/csv/validate-csv ,

error message:

"Error 1 title: CSV Contains Empty Lines message: The file contains empty lines. row: 2654"

Missing value due to formatting issues while editing the file, found the distinct values on the file and replaced the empty cell with the only missing state for that time frame.

<!-- grade_id: cell-e72d44bc996d8aaf, points: 0, grade: false, solution: true -->
```python
import requests
import csv
import json
import xmltodict

def parseFile(datadict, format):

    # create important variables 
    dataformat = datadict["detectedFormat"]
    response = requests.get(datadict["resourceURL"])
    
    # VALITATE CSV CONTENT
    if format == dataformat == "CSV":
        
        # list of delimiter to be checked
        delimiter = [",",";","\t"]
        # gets the text of the URL and dumps each line into a list
        text = response.text.splitlines()
        best_delimiter = None
        
        for d in delimiter:
            # tries if any of the values of delimiter are eligable for a delimiter
            try:
                # reads the text of the file with a possible delimiter of the delimiter list 
                reader = csv.reader(text, delimiter = d)
    
                for row in reader:
                    # create a set with all the row lengths. used characteristic that the set has no duplicates (len(set) = 1 if all rows are same length)
                    row_len = {len(row)}

                # find best dilimiter
                # 1. condition checks if all row lengths are the same
                # 2. condition checks if there are multiple colums in the csv file (transform set into list to access through indices)
                if len(row_len) == 1 and list(row_len)[0] > 1: 
                    best_delimiter = d                
            except Exception as e:
                print(e)
            
        if best_delimiter:
            csv_reader = csv.reader(text, delimiter = best_delimiter)
                
            return True
        
    
    # VALITATE JSON CONTENT
    elif format == dataformat == "JSON":
        # try if content of the dataset could be parsed as a json file
        try: 
            data_json = response.json()
           
            return True
        except Exception as e:
            
            return False
        
    # VALITATE XML CONTENT
    elif format == dataformat == "XML":
        # try if content of the dataset could be parsed as a xml file        
        try:
            data_xml = xmltodict.parse(response.content)
            
            return True
        except Exception as e:
            
            return False
    else:
        return False
    raise NotImplementedError()
```

<!-- grade_id: cell-44a0730c406d4f1f, points: 0.375, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_in, assert_true
assert_equal([parseFile(dataset1, "XML"),
    parseFile(dataset1, "JSON"),
    parseFile(dataset1, "CSV"),
    parseFile(dataset2, "XML"),
    parseFile(dataset2, "JSON"),
    parseFile(dataset2, "CSV")].count(True), 2)
```

<!-- grade_id: cell-a6fe57057ce3e18e, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-aa7fd9b40f36f086, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-7d845903b5b4589e, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-d647097f451b66c2, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-f77b7d77dfebf6fc, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-f7ec87ba2b1d4805, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-0872aa870b41a0f9, points: 0.375, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

-----
## Step 3 - Content analysis (5 points)

Similar to the Python function `parseFile(datadict,format)` above, now create a new Python function `describeFile(datadict)` that analyses the given file according to the respective format detected in Step 1 and returns a dictionary containing the following information:

* for CSV files: number of columns, number of rows, column number (from 0 to n) of the column which contains the longest text. You do not have to try to transform any string to integer or float, simply take the values as is from the csv file. That is, the resulting dictionary should have the following form:

    ```
    { "numberOfColumns:"  ...,
       "numberOfRows":  ... ,
       "longestColumn" : ... }
    ```

* for JSON files: number of different attribute names, nesting depth, length of the longest list appearing in an attribute value. That is, the resulting dictionary should have the following form:

    ```
    { "numberOfAttributes:" ... ,
      "nestingDepth":  ... ,
      "longestListLength" : ... }
     ```

  Here the `longestListLength` should be set to 0 if no list appears. [Nesting depth](https://www.tutorialspoint.com/find-depth-of-a-dictionary-in-python) is defined as follows: 
   * a flat JSON object with only atomic attribute values has depth 1. 
   * a JSON attribute with another object as value (or another oject as member of a list value!) increases the depth by 1
   * and so on.


* for XML files: number of different element and attribute a names (i.e. the sum of both), nesting depth, maximum numeric value in the dataset. That is, the resulting dictionary should have the following form:

    ```
    { "numberOfElementsAttributes:" ... ,
      "nestingDepth":  ... ,
      "maxNumericValue" : ... }
     ```

  Here the `maxNumericValue` should be set to 0 if there are no numberic values present. Nesting depth is defined as the nesting depth of elements.
  
For files that cannot be parsed with respective given format, the function should simply return an empty dictionary (`{}`).

<!-- grade_id: cell-223d8521820ec726, points: 0, grade: false, solution: true -->
```python
import codecs

def describeCSV(datadict):
    URL = datadict["resourceURL"]
    response = requests.get(URL)
    if response.status_code == 200:
    # Decode content to text
        csv_content = response.content.decode('utf-8')
        csv_reader = csv.reader(csv_content.splitlines(), delimiter=',')
        
        num_rows = 0
        num_columns = 0
        rows = []  # dont print this - too big, crashed my jupyter 
    
        for row in csv_reader:
            rows.append(row)
            num_rows += 1
            
        if rows:
            num_columns = len(rows[0])
        
        max_len = ""
        col_of_max_len = 0

        for x in rows:
            for col_index, value in enumerate(x):
                if len(value) > len(max_len):
                    max_len = value
                    col_of_max_len = col_index
                   
             
        return {
            "numberOfColumns:": num_columns,
            "numberOfRows:":num_rows,
            "longestColumn:": col_of_max_len,
        }
    
    else:
        print("Failed CSV Download:", response.status_code)

def finddepthJSON(datadict): ### https://www.tutorialspoint.com/find-depth-of-a-dictionary-in-python, 18.04.2024
    if isinstance(datadict, dict):
        return 1 + (max(map(finddepthJSON, datadict.values()))
            if datadict else 0)
    return 0

def JsonLongestList(datadict):
    longestList = 0
    
    for x in datadict:
        if isinstance(datadict, list):
            if len(x) > longestList:
                longestList = len(x)
    return longestList

def describeJSON(datadict):
    URL = datadict["resourceURL"]
    response = requests.get(URL)
    if response.status_code == 200:
        json_data = response.json()
        
        
        CountAttributes = len(json_data[0]) ## Counting the attributes in the datafile
       
        return {
            "numberOfAttributes:": CountAttributes,
            "nestingDepth:": finddepthJSON(json_data[0]),
            "longestListLength": JsonLongestList(json_data[0])
        }
     
    else:
            print("Failed JSON Download:", response.status_code)

def finddepthXML(datadict): ### https://www.tutorialspoint.com/find-depth-of-a-dictionary-in-python, 18.04.2024
    if isinstance(datadict, dict):
        return 1 + (max(map(finddepthXML, datadict.values()))
            if datadict else 0)
    return 0

def maxValue(a):
    max_value = 0
    
    def recursion(obj):
        # access non local variables to modify them within the new function
        nonlocal max_value
        if isinstance(obj, dict):
            for key, value in obj.items():
                recursion(value)
        if isinstance(obj, list):
            for element in obj:
                recursion(element)
        value = 0
        try:
            value = float(obj)
        except:
            pass
 
        # adds every integer to the counter
        if isinstance(value, (int, float)) and value > max_value:
            max_value = value

    recursion(a)
    result = (max_value)
    return(result)

def Count_Attributes_Elements(datadict):
    elements_attributes_Counter = 0
    if isinstance(datadict, dict):
        elements_attributes_Counter += len(datadict)
        for key, value in datadict.items():
            if isinstance(value, dict):                   
                elements_attributes_Counter += Count_Attributes_Elements(value)
            elif isinstance(value, list):
                for i in value:
                    elements_attributes_Counter += Count_Attributes_Elements(i)
                    Count_Attributes_Elements(value)
        return elements_attributes_Counter            

def describeXML(datadict):
    URL = datadict["resourceURL"]
    response = requests.get(URL)
    if response.status_code == 200:
        xml_data = response.text
        xml_dict = xmltodict.parse(xml_data)
        
        return {
            "numberOfAttributes:":Count_Attributes_Elements(xml_dict),
            "nestingDepth:": finddepthXML(xml_dict),
            "maxNumericValue": maxValue(xml_dict)
        }
       
    else:
        print("Failed XML Download")
        return none            

def describeFile(datadict):
    if datadict["detectedFormat"] == "CSV":
        return describeCSV(datadict)
    elif datadict["detectedFormat"] == "JSON":
        return describeJSON(datadict)
    elif datadict["detectedFormat"] == "XML":
        return describeXML(datadict)
    else:
        return {}
    raise NotImplementedError()
```

<!-- grade_id: cell-eba7f348525a45dd, points: 0.1875, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_in, assert_true
assert_equal(len(describeFile(dataset1)), 3)
assert_equal(len(describeFile(dataset2)), 3)
```

<!-- grade_id: cell-557eee89020c4510, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-b125d2724a7003c1, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-c366624cc0d29728, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-2644ceec88376e7e, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-75825a6fa21b6764, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-1c787c4413548bab, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-2e575d64dfdbaa9b, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-4218dd06c8ea8833, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-91e3b4cbf0931744, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-6fb522b91576b08f, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-5bab3319a1f792f9, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-2213630186aef937, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-136ca0d44e2fd03d, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-024d562b2d63af48, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

<!-- grade_id: cell-469f1889794c4746, points: 0.1875, grade: true, solution: false -->
```python
### DO NOT DELETE OR CHANGE THIS CELL!
```

***Final check:***

Be sure to cross-check the results by
1) manually inspecting your chosen dataset and
2) comparing the results for plausibility against the results of your code... 

Describe your findings and use the following structure for your answer below:

*Data set 1*

(number and types of items etc., describe your findings)

*Data set 2*

(number and types of items etc., describe your findings)

<!-- grade_id: cell-4709b3248c430b6a, points: 2, grade: true, solution: true -->
Data set 1

{'numberOfAttributes:': 7, 'nestingDepth:': 1, 'longestListLength': 0}
These results are correct for our Files and seem to be correct for a few other Testfiles

Data set 2

{'numberOfColumns:': 10, 'numberOfRows:': 2653, 'longestColumn:': 0}
These results are correct for our Files and seem to be correct for a few other Testfiles